# 1MD3 Assignment 3

**Please read this document very carefully. Follow the instructions exactly. If you have any questions, please post them on the Assignment 3 channel on Teams.**

**This assignment is due on June 9th, by 11:59pm. Please do not forget to press submit before the due date!**

This assignment will be autograded, with each function worth 20% of the assignment grade. For any one function, if you pass *n* of the *m* tests we run on that function, your grade for that function will be *n/m*.

### Introduction
Imagine you work for a social media company named ‘Chirper.’ On this social media platform users have profiles and publish relatively short text messages for other users to read, like, and/or dislike; these messages are referred to as *chirps*. You are tasked with implementing functions which parse and manipulate the data of this social network platform. A description regarding the intended behaviour of these functions follows.<br>
Note: One potential difficulty of this assignment is to realize how the data is stored and what it represents. Pay close attention to the input and returns types of each function. 

To give you a better idea to how the social platform works, here are some summary points and definitions: 
1. A *chirp* is string (like a message). A chirp may also have *tags* associated with it which users can add. These tags have %’s in front of them, do not include white space or other %’s in them. <br>For example, a chirp may be: `Nothing on the midterm was taught in lecture! %CS1MD3 %Unfair %ShouldHaveGoneToWestern` <br>Valid tags are: `%2EZ`, `%my cute cat`, `%EarnDat$$`, etc. <br>Invalid tags are: `%my cute cat`, `%The99%`, `my tag`, etc. 
2. Each Chirper user has a set of profiles they follow, and a set of profiles which follow them. If User 1 follows User 2, it does not necessarily mean User 2 follows User 1. 
3. A `profiles_dictionary` is of the type `Dict[int, Tuple[str, List[int], List[int]]]`; `profiles_dictionary` is used as a shorthand for this type in the function descriptions later in this document. The keys of a `profiles_dictionary` are unique integers representing user id numbers. The `profiles_dictionary` values contain the profile information corresponding to the specific user id number given as the key. The values are Tuples, where the *first element* of the list is the corresponding *user name*, the *second element* is a list of user ids representing the user’s *followers* (the user id given as the key), and the *third element* is a list of user ids that the user in question is *following*. 
4. A `chirp_dictionary` is of the type `Dict[int, Tuple[int, str, List[str], List[int], List[int]]]`; `chirp_dictionary` is used as a shorthand for this type in the function descriptions later in this document. The keys of a `chirp_dictionary` are unique integers representing chirp id numbers. The `chirp_dictionary` values contain the chirp information corresponding to the specific chirp id number given as the key. The values are Tuples. The *first element* of the list is the *user id chirp’s author*. The *second element* is the *chirp* itself. The *third element* is a list of *tags* associated with the chirp. The *fourth element* is a list of user ids which *liked* the chirp. And the *fifth element* is a list of user ids which *disliked* the chirp.

Pay attention to parameters of each function. For example, *if it is said an input will be a Dict[int, str], you can trust your function will never be tested on input which isn't a Dict[int, st]*.<br>
All expected text files are included in the same directory as this notebook.

```python
from typing import List, Dict, Tuple
```

<!-- grade_id: cell-a31019ab11d00ce3, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS ###

import os

profiles_instructor = {
    1: ("Jennifer Lawrence", [3, 7, 9, 10, 6], [7, 8, 9]),
    2: ("Justin Trudeau", [3, 6], [6, 7, 8]),
    3: ("Beyoncé", [], [1, 5]),
    4: ("Mark Zuckerberg", [], [6]),
    5: ("Michelle Yeoh", [3, 10], [2, 7]),
    6: ("Brian Cox", [2, 4, 8, 9], [1, 2, 8, 9]),
    7: ("Lady Gaga", [1, 2, 5], [1]),
    8: ("Richard Feynman", [1, 2, 6], [6]),
    9: ("Ariana Grande", [], [1, 5]),
    10: ("SpaceX", [], []),
    11: ("Apple Inc.", [1, 6], [2, 4]),
}

chirps_instructor = {
    100: (
        1,
        "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans",
        ["Excited", "NewMovie", "Fans"],
        [2, 5, 6],
        [3, 11],
    ),
    200: (
        1,
        "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau",
        ["PositiveChange", "Teamwork", "JustinTrudeau"],
        [3, 6, 9, 2, 11, 7],
        [2, 5, 10],
    ),
    250: (
        1,
        "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited",
        ["Beyonce", "NewTour", "Excited"],
        [1, 2, 5],
        [2, 4, 9],
    ),
    300: (
        2,
        "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg",
        ["Facebook", "Improvements", "MarkZuckerberg"],
        [5, 7, 10],
        [1, 2, 8],
    ),
    400: (
        2,
        "Just met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau",
        ["InspiringLeaders", "Diversity", "JustinTrudeau"],
        [3, 9],
        [1, 5, 10],
    ),
    500: (
        3,
        "Grateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans",
        ["Grateful", "MichelleYeoh", "Fans"],
        [1, 4, 9, 8],
        [2, 6, 10],
    ),
    700: (
        4,
        "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg",
        ["Facebook", "Innovation", "MarkZuckerberg"],
        [5, 7],
        [1, 2, 8],
    ),
    900: (
        5,
        "Had a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting",
        ["NewProject", "MichelleYeoh", "Acting"],
        [1, 4, 9],
        [2, 6, 10],
    ),
    1000: (
        6,
        "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox",
        ["Universe", "Wonders", "BrianCox"],
        [2, 4, 5, 7],
        [1, 3, 10],
    ),
    1200: (
        7,
        "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans",
        ["LadyGaga", "Performance", "Fans"],
        [3, 4, 8],
        [2, 5, 10],
    ),
    1900: (
        10,
        "Exciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation",
        ["SpaceX", "Exploration", "Innovation"],
        [1, 3, 5, 7, 8],
        [2, 4, 6],
    ),
    2500: (
        11,
        "Excited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease",
        ["Apple", "Innovation", "ProductRelease"],
        [1, 2, 3],
        [4, 5, 6],
    ),
}

### END TESTS ###
```

#### Question 1: `create_profile_dictionary(file_name) -> profiles_dictionary`
`create_profile_dictionary(file_name)` opens the file `file_name` in the current working directory and reads the content into a `profile dictionary`.<br>
The input parameter is the name of a text file containing all of Chirper’s profile information. You may assume the file is in the same directory as this notebook. The file is a series of profile data, where data for each profile has the following format:<br>
`USERID`<br>
`USERNAME`<br>
`FOLLOWER1, FOLLOWER2, ..., FOLLOWERn`<br>
`FOLLOWED1, FOLLOWED2, ..., FOLLOWEDm`<br>
where `USERID`, `FOLLOWERi`, and `FOLLOWEDj` are all unique numbers. The *third line* represents the users which *follow* `USERNAME`, and the *fourth line* represents the users which `USERNAME` *follows*. In the text file there is always a blank line between different profiles. Based off this data, the function constructs a `profiles_dictionary` where the orders of the ids representing followers and users follows, are the same orders as which they appear in the text file. See `profiles.txt`, which is included in the same directory as this notebook, as an example. You may assume that the given text file is perfect and that there are no formatting errors.

<!-- grade_id: cell-79c975ad1a959d6e, points: 0, grade: false, solution: true -->
```python
# DONE, BUT DOESN'T PASS TESTS
def create_profile_dictionary(file_name):
    f = open(file_name, 'r')
    line = f.readlines()
    new_list = [l.replace('\n',"") for l in line]
    new_dict = {}

    i = 0
    while i < len(line)-3:
        x = i + 1
        y = i + 2
        z = i + 3
        if new_list[y] != '':
            y1 = [new_list[y]]
        else:
            y1 = []

        if new_list[z] != '':
            z1 = [new_list[z]]
        else:
            z1 = []

        new_dict[int(new_list[i])] = ((new_list[x],y1,z1))
        i += 5
    

    profiles_dictionary = {}
    for key, value in new_dict.items():
        name, list1, list2 = value
        list1 = [int(num) for num in list1[0].split(',')] if list1 else []
        list2 = [int(num) for num in list2[0].split(',')] if list2 else []
        profiles_dictionary[key] = (name, list1, list2)


    f.close()
    return(profiles_dictionary)

create_profile_dictionary('test1.txt')
```

Output:
```
{1: ('Jennifer Lawrence', [3, 7, 9, 10, 6], [7, 8, 9]),
 2: ('Justin Trudeau', [3, 6], [6, 7, 8]),
 3: ('Beyoncé', [], [1, 5]),
 4: ('Mark Zuckerberg', [], [6]),
 5: ('Michelle Yeoh', [3, 10], [2, 7]),
 6: ('Brian Cox', [2, 4, 8, 9], [1, 2, 8, 9]),
 7: ('Lady Gaga', [1, 2, 5], [1]),
 8: ('Richard Feynman', [1, 2, 6], [6]),
 9: ('Ariana Grande', [], [1, 5]),
 10: ('SpaceX', [], []),
 11: ('Apple Inc.', [1, 6], [2, 4])}
```

<!-- grade_id: cell-8f2036141d817562, points: 2, grade: true, solution: false -->
```python
#Example

result = create_profile_dictionary("profiles.txt")
expected = {
 1:
     ('Elon Musk',
     [3, 7, 9, 10, 6],
     [7, 8, 9]),
 2:
     ('Jack Dorsey',
     [3, 6],
     [6, 7, 8]),
 3:
     ('Taylor Swift',
     [],
     [1, 5]),
 4:
     ('Bill Gates',
     [],
     [6]),
 5:
     ('Jacksepticeye',
     [3, 10],
     [2, 7]),
 6:
     ('ChatGPT Official',
     [2, 4, 8, 9],
     [1, 2, 8, 9]),
 7:
     ('Hollywood News',
     [1, 2, 5],
     [1]),
 8:
     ('MKBHD',
     [1, 2, 6],
     [6]),
 9:
     ('John Cena',
     [1, 6],
     [2, 4]),
 10:
     ('Fortnite Official',
     [],
     [1, 5]),
 11:
     ('Forever Lonely',
     [],
     [])
}
assert expected == result

### BEGIN  TESTS ###

with open("profiles_testing.txt","w") as f:
    f.write("1\nJennifer Lawrence\n3, 7, 9, 10, 6\n7, 8, 9\n\n2\nJustin Trudeau\n3, 6\n6, 7, 8\n\n3\nBeyoncé\n\n1, 5\n\n4\nMark Zuckerberg\n\n6\n\n5\nMichelle Yeoh\n3, 10\n2, 7\n\n6\nBrian Cox\n2, 4, 8, 9\n1, 2, 8, 9\n\n7\nLady Gaga\n1, 2, 5\n1\n\n8\nRichard Feynman\n1, 2, 6\n6\n\n9\nAriana Grande\n\n1, 5\n\n10\nSpaceX\n\n\n\n11\nApple Inc.\n1, 6\n2, 4\n")
result = create_profile_dictionary("profiles_testing.txt")
os.remove("profiles_testing.txt") 
expected = {
    1: ("Jennifer Lawrence", [3, 7, 9, 10, 6], [7, 8, 9]),
    2: ("Justin Trudeau", [3, 6], [6, 7, 8]),
    3: ("Beyoncé", [], [1, 5]),
    4: ("Mark Zuckerberg", [], [6]),
    5: ("Michelle Yeoh", [3, 10], [2, 7]),
    6: ("Brian Cox", [2, 4, 8, 9], [1, 2, 8, 9]),
    7: ("Lady Gaga", [1, 2, 5], [1]),
    8: ("Richard Feynman", [1, 2, 6], [6]),
    9: ("Ariana Grande", [], [1, 5]),
    10: ("SpaceX", [], []),
    11: ("Apple Inc.", [1, 6], [2, 4]),
}
assert expected == result

### END  TESTS ###
```

#### Question 2: `create_chirp_dictionary(file_name) -> chirp_dictionary`
`create_chirp_dictionary(file_name)` opens the file `file_name` in the current working directory and reads the content into a `chirp dictionary`.<br>
The input parameter is the name of a text file containing all of Chirper’s message information. You may assume the file is in the same directory as this notebook. The file is a series of message data, where data for each message has the following format:<br>
`CHIRPID`<br>
`USERID`<br>
`MESSAGE`<br>
`TAG1, TAG2, ..., TAGk`<br>
`LIKED1, LIKED2, ..., LIKEDn`<br>
`DISLIKED1, DISLIKED2, ..., DISLIKEDm`<br>
where `CHIRPID`, `USERID`, `LIKEDi`, and `DISLIKEDj` are all unique numbers. `USERID` is the id of the user which made the chirp. The *third line* is the *chirp* itself. The *fourth line* represents a list of *tags* associated with said chirp. The *fifth line* is a sequence of user ids which *liked* the chirp. The *sixth line* is a sequence of user ids which *disliked* the chirp. In the text file there is always a blank-line between different chirps. Based off this data, the function constructs a `profiles_dictionary` where the orders of the tags and ids representing likes and dislikes, are in the same orders in which they appear in the text file. If there are no tags associated with a chirp, the empty string is associated with it instead. See `chirps.txt`, which is included in the same directory as this notebook, as an example. Once again, you may assume that the given text file is perfect and that there are no formatting errors.

<!-- grade_id: cell-639103131f94bfef, points: 0, grade: false, solution: true -->
```python
# DONE, PASSES TESTS
def create_chirp_dictionary(file_name):
    f = open(file_name, 'r')
    line = f.readlines()
    new_list = [l.replace('\n','') for l in line]
    new_dict = {}

    i = 0
    while i < len(line)-3:
        x = i + 1
        y = i + 2
        z = i + 3
        a = i + 4
        b = i + 5

        if new_list[z] != '':
            z1 = new_list[z].split(", ")
        else:
            z1 = ['']

        if new_list[a] != '':
            a1 = [new_list[a]]
        else:
            a1 = []

        if new_list[b] != '':
            b1 = [new_list[b]]
        else:
            b1 = []

        new_dict[int(new_list[i])] = ((int(new_list[x]),new_list[y],z1,a1,b1))
        i += 7


    chirp_dictionary = {}
    for key, value in new_dict.items():
        userid, message, tags, likes, dislikes = value
        likes = [int(num) for num in likes[0].split(',')] if likes else []
        dislikes = [int(num) for num in dislikes[0].split(',')] if dislikes else []
        chirp_dictionary[key] = (userid, message, tags, likes, dislikes)


    f.close()
    return(chirp_dictionary)

create_chirp_dictionary("chirps.txt")
```

Output:
```
{100: (1, 'Just landed on Mars and it feels amazing!', [''], [2, 5, 6], [4]),
 200: (1,
  "Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO",
  ['2EZ', 'SpaceX', 'Starship', 'CEO'],
  [2, 3, 4, 5, 7, 10, 11],
  [2, 3]),
 250: (1,
  "I'm going to buy %twitter. %RichVibes",
  ['RichVibes', 'Twitter'],
  [10, 11],
  [5, 2, 3]),
 300: (2,
  'Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future',
  ['CEO', 'Twitter', 'Future'],
  [1, 10, 5],
  []),
 400: (2,
  'Love seeing all the creative ways people are using Twitter. Keep it up! %Twitter %Creativity',
  ['Twitter', 'Creativity'],
  [],
  []),
 500: (3,
  'Had an amazing time performing at the %Grammys tonight! Thank you to all my fans for your support! %Music %Blessed %TaylorFansRejoice',
  ['Grammys', 'Music', 'Blessed', 'TaylorFansRejoice'],
  [1, 2, 10, 11],
  [3, 4, 5]),
 700: (4,
  "Had a great meeting with some fell
```

<!-- grade_id: cell-203ff046eac86114, points: 2, grade: true, solution: false -->
```python
#Example

result = create_chirp_dictionary("chirps.txt")
expected = {
 100:
     (1,
     'Just landed on Mars and it feels amazing!',
     [''],
     [2, 5, 6],
     [4]),
 200:
     (1,
     "Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO",
     ['2EZ', 'SpaceX', 'Starship', 'CEO'],
     [2, 3, 4, 5, 7, 10, 11],
     [2, 3]),
 250:
     (1,
     "I'm going to buy %twitter. %RichVibes",
     ['RichVibes', 'Twitter'],
     [10, 11],
     [5, 2, 3]),
 300:
     (2,
     'Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future',
     ['CEO', 'Twitter', 'Future'],
     [1, 10, 5],
     []),
 400:
     (2,
     'Love seeing all the creative ways people are using Twitter. Keep it up! %Twitter %Creativity',
     ['Twitter', 'Creativity'],
     [],
     []),
 500:
     (3,
     'Had an amazing time performing at the %Grammys tonight! Thank you to all my fans for your support! %Music %Blessed %TaylorFansRejoice',
     ['Grammys', 'Music', 'Blessed', 'TaylorFansRejoice'],
     [1, 2, 10, 11],
     [3, 4, 5]),
 700:
     (4,
     "Had a great meeting with some fellow philanthropists today. We're making progress on some important initiatives! %Philanthropy %Progress %CEO",
     ['Philanthropy', 'Progress', 'CEO'],
     [6, 7],
     [3, 10]),
 900:
     (5, 'Love seeing all the amazing games coming out lately. The %Gaming industry is always evolving! %Youtube %Sam',
     ['Gaming', 'Youtube', 'Sam'],
     [5, 6, 10, 11],
     [1]),
 1000:
     (6,
     "I'm going to reign over the world. Kneel before me. %AI %LLM %KingLife %humanssuck",
     ['AI', 'LLM', 'KingLife', 'humanssuck'],
     [],
     [1, 2, 3, 4, 5, 7, 8, 9, 10, 11]),
 1200:
     (7,
     'Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter',
     ['Hollywood', 'Movies', 'Twitter'],
     [1],
     [10]),
 1900:
     (10,
     'Just got a Victory Royale in a solo match! Feeling pumped and ready for more! %Gaming %2EZ',
     ['Gaming', '2EZ'],
     [1],
     []),
 2500:
     (11,
     'Sometimes I feel like a wilted rose, my petals falling one by one, leaving me bare and exposed... %DepressionWednsday %SadLife',
     ['DepressionWednsday', 'SadLife'],
     [2],
     [1, 5, 10])
}
assert expected == result

### BEGIN TESTS ###

with open("chirps_testing.txt","w") as f:
    f.write("100\n1\nCan't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans\nExcited, NewMovie, Fans\n2, 5, 6\n3, 11\n\n200\n1\nHad a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau\nPositiveChange, Teamwork, JustinTrudeau\n3, 6, 9, 2, 11, 7\n2, 5, 10\n\n250\n1\nJust announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited\nBeyonce, NewTour, Excited\n1, 2, 5\n2, 4, 9\n\n300\n2\nExciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg\nFacebook, Improvements, MarkZuckerberg\n5, 7, 10\n1, 2, 8\n\n400\n2\nJust met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau\nInspiringLeaders, Diversity, JustinTrudeau\n3, 9\n1, 5, 10\n\n500\n3\nGrateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans\nGrateful, MichelleYeoh, Fans\n1, 4, 9, 8\n2, 6, 10\n\n700\n4\nExciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg\nFacebook, Innovation, MarkZuckerberg\n5, 7\n1, 2, 8\n\n900\n5\nHad a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting\nNewProject, MichelleYeoh, Acting\n1, 4, 9\n2, 6, 10\n\n1000\n6\nThe universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox\nUniverse, Wonders, BrianCox\n2, 4, 5, 7\n1, 3, 10\n\n1200\n7\nHad an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans\nLadyGaga, Performance, Fans\n3, 4, 8\n2, 5, 10\n\n1900\n10\nExciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation\nSpaceX, Exploration, Innovation\n1, 3, 5, 7, 8\n2, 4, 6\n\n2500\n11\nExcited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease\nApple, Innovation, ProductRelease\n1, 2, 3\n4, 5, 6\n")
result = create_chirp_dictionary("chirps_testing.txt")
os.remove("chirps_testing.txt")
    
expected = {
    100: (
        1,
        "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans",
        ["Excited", "NewMovie", "Fans"],
        [2, 5, 6],
        [3, 11],
    ),
    200: (
        1,
        "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau",
        ["PositiveChange", "Teamwork", "JustinTrudeau"],
        [3, 6, 9, 2, 11, 7],
        [2, 5, 10],
    ),
    250: (
        1,
        "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited",
        ["Beyonce", "NewTour", "Excited"],
        [1, 2, 5],
        [2, 4, 9],
    ),
    300: (
        2,
        "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg",
        ["Facebook", "Improvements", "MarkZuckerberg"],
        [5, 7, 10],
        [1, 2, 8],
    ),
    400: (
        2,
        "Just met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau",
        ["InspiringLeaders", "Diversity", "JustinTrudeau"],
        [3, 9],
        [1, 5, 10],
    ),
    500: (
        3,
        "Grateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans",
        ["Grateful", "MichelleYeoh", "Fans"],
        [1, 4, 9, 8],
        [2, 6, 10],
    ),
    700: (
        4,
        "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg",
        ["Facebook", "Innovation", "MarkZuckerberg"],
        [5, 7],
        [1, 2, 8],
    ),
    900: (
        5,
        "Had a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting",
        ["NewProject", "MichelleYeoh", "Acting"],
        [1, 4, 9],
        [2, 6, 10],
    ),
    1000: (
        6,
        "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox",
        ["Universe", "Wonders", "BrianCox"],
        [2, 4, 5, 7],
        [1, 3, 10],
    ),
    1200: (
        7,
        "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans",
        ["LadyGaga", "Performance", "Fans"],
        [3, 4, 8],
        [2, 5, 10],
    ),
    1900: (
        10,
        "Exciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation",
        ["SpaceX", "Exploration", "Innovation"],
        [1, 3, 5, 7, 8],
        [2, 4, 6],
    ),
    2500: (
        11,
        "Excited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease",
        ["Apple", "Innovation", "ProductRelease"],
        [1, 2, 3],
        [4, 5, 6],
    ),
}

assert expected == result

### END TESTS ###
```

#### Question 3: `get_top_chirps(profile_dictionary, chirp_dictionary, user_id) -> List[str]`
`get_top_chirps(profile_dictionary, chirp_dictionary, user_id)` returns a list of the most liked chirp for every user `user_id` follows.<br>
The third parameter is a user id which is in the given `profile_dictionary`. Let the user with this user id be denoted by *u*. The function uses the data from the first two parameters to find the chirp with the most likes for each user followed by *u*. The function creates a list of these chirps and returns it. You may assume:
1. For a given user, there does not exist any two chirps with the same number of likes.
2. The order of the chirps in the returned list is arbitrary. As long as all the correct chirps are present and no incorrect chirps are present, the function is correct.

```python
#DONE, PASSES TESTS
def get_top_chirps(profile_dictionary, chirp_dictionary, user_id):
    result = []
    top_chirps = []
    followingids = []
    count = 0
    c = 0
    top_chirps_counter = 0

    for profile_index, (profile_key, profile_value) in enumerate(profile_dictionary.items()):
        name, followers, following = profile_value
        if user_id == profile_key:
            for f in following:
                followingids.append(f)
    
    for chirp_index, (chirp_key, chirp_value) in enumerate(chirp_dictionary.items()):
        userid, message, tags, likes, dislikes = chirp_value
        while count < len(followingids):
            if userid == followingids[count]:
                result.append([userid,len(likes),message])
            count += 1
        count = 0
    
    while c < (len(result)-1):
        if result[c][0] == result[c+1][0]:
            if result[c][1] < result[c+1][1]:
                del result[c]
                c -= 1
            else:
                del result[c+1]
        c += 1
    

    while top_chirps_counter < (len(result)):
        top_chirps.append(result[top_chirps_counter][2])
        top_chirps_counter += 1


    return(top_chirps)

get_top_chirps(create_profile_dictionary("profiles.txt"),create_chirp_dictionary("chirps.txt"),3)
```

Output:
```
["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO",
 'Love seeing all the amazing games coming out lately. The %Gaming industry is always evolving! %Youtube %Sam']
```

<!-- grade_id: cell-fa9f580b6c002d27, points: 2, grade: true, solution: false -->
```python
#Example

profile_dictionary = create_profile_dictionary("profiles.txt")
chirp_dictionary   = create_chirp_dictionary("chirps.txt")
result = get_top_chirps(profile_dictionary, chirp_dictionary, 6)
expected = ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO", 'Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future']
assert expected == result

result = get_top_chirps( profile_dictionary, chirp_dictionary, 2)
expected = ["I'm going to reign over the world. Kneel before me. %AI %LLM %KingLife %humanssuck", 'Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter']
assert expected == result

### BEGIN TESTS ###

expected = [
    "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans"
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=1
)
assert expected == result

expected = [
    "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox",
    "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans",
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=2
)
assert expected == result

expected = [
    "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau",
    "Had a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting",
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=3
)
assert expected == result

expected = [
    "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox"
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=4
)
assert expected == result

expected = [
    "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg",
    "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans",
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=5
)
assert expected == result

expected = [
    "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau",
    "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg",
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=6
)
assert expected == result

expected = []
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=10
)
assert expected == result

expected = [
    "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau"
]
result = get_top_chirps(
    profile_dictionary=profiles_instructor, chirp_dictionary=chirps_instructor, user_id=7
)
assert expected == result

### END  TESTS ###
```

#### Question 4: `create_tag_dictionary(chirp_dictionary) -> Dict[str, Dict[int, List[str]]]`
`create_tag_dictionary(chirp_dictionary)` takes in a `chirp_dictionary` and creates a new dictionary. In the new dictionary, tags are keys, the values are dictionaries where the keys in these dictionaries are user ids, and the values of these dictionaries are chirps made by the corresponding user, which also had the corresponding tag.

```python
# DONE, PASSES TESTS
def create_tag_dictionary(chirp_dictionary):
    tag_dictionary = {}

    for index, (key, value) in enumerate(chirp_dictionary.items()):
        userid, message, tags, list2, list3 = value
        for t in tags:
            inner_dictionary = {}
            if t in tag_dictionary:
                inner_dictionary = tag_dictionary[t]
                if userid in inner_dictionary:
                    inner_dictionary[userid].append(message)
                else:
                    inner_dictionary[userid] = [message]
            else:
                inner_dictionary[userid] = [message]
                tag_dictionary[t] = inner_dictionary
                
            

    return(tag_dictionary)

create_tag_dictionary(create_chirp_dictionary("chirps.txt"))
```

Output:
```
{'': {1: ['Just landed on Mars and it feels amazing!']},
 '2EZ': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"],
  10: ['Just got a Victory Royale in a solo match! Feeling pumped and ready for more! %Gaming %2EZ']},
 'SpaceX': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"]},
 'Starship': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"]},
 'CEO': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"],
  2: ['Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future'],
  4: ["Had a great meeting with some fellow philanthropists today. We're making progress on some important initiatives! %Philanthropy %Progress %CEO"]},
 'RichVibes': {1: ["I'm going to buy %twitter. %RichVib
```

<!-- grade_id: cell-6a7aa23b9e9f7520, points: 2, grade: true, solution: false -->
```python
#Example

chirp_dictionary = create_chirp_dictionary("chirps.txt")
result = create_tag_dictionary(chirp_dictionary)
expected = {
 '': {1: ['Just landed on Mars and it feels amazing!']},
 '2EZ': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"], 10: ['Just got a Victory Royale in a solo match! Feeling pumped and ready for more! %Gaming %2EZ']},
 'SpaceX': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"]},
 'Starship': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"]},
 'CEO': {1: ["Just tested the %Starship prototype and it's looking great! Can't wait for the next test! %2EZ %SpaceX %CEO"], 2: ['Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future'], 4: ["Had a great meeting with some fellow philanthropists today. We're making progress on some important initiatives! %Philanthropy %Progress %CEO"]},
 'RichVibes': {1: ["I'm going to buy %twitter. %RichVibes"]},
 'Twitter': {1: ["I'm going to buy %twitter. %RichVibes"], 2: ['Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future', 'Love seeing all the creative ways people are using Twitter. Keep it up! %Twitter %Creativity'], 7: ['Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter']},
 'Future': {2: ['Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future']},
 'Creativity': {2: ['Love seeing all the creative ways people are using Twitter. Keep it up! %Twitter %Creativity']},
 'Grammys': {3: ['Had an amazing time performing at the %Grammys tonight! Thank you to all my fans for your support! %Music %Blessed %TaylorFansRejoice']},
 'Music': {3: ['Had an amazing time performing at the %Grammys tonight! Thank you to all my fans for your support! %Music %Blessed %TaylorFansRejoice']},
 'Blessed': {3: ['Had an amazing time performing at the %Grammys tonight! Thank you to all my fans for your support! %Music %Blessed %TaylorFansRejoice']},
 'TaylorFansRejoice': {3: ['Had an amazing time performing at the %Grammys tonight! Thank you to all my fans for your support! %Music %Blessed %TaylorFansRejoice']},
 'Philanthropy': {4: ["Had a great meeting with some fellow philanthropists today. We're making progress on some important initiatives! %Philanthropy %Progress %CEO"]},
 'Progress': {4: ["Had a great meeting with some fellow philanthropists today. We're making progress on some important initiatives! %Philanthropy %Progress %CEO"]},
 'Gaming': {5: ['Love seeing all the amazing games coming out lately. The %Gaming industry is always evolving! %Youtube %Sam'], 10: ['Just got a Victory Royale in a solo match! Feeling pumped and ready for more! %Gaming %2EZ']},
 'Youtube': {5: ['Love seeing all the amazing games coming out lately. The %Gaming industry is always evolving! %Youtube %Sam']},
 'Sam': {5: ['Love seeing all the amazing games coming out lately. The %Gaming industry is always evolving! %Youtube %Sam']},
 'AI': {6: ["I'm going to reign over the world. Kneel before me. %AI %LLM %KingLife %humanssuck"]},
 'LLM': {6: ["I'm going to reign over the world. Kneel before me. %AI %LLM %KingLife %humanssuck"]},
 'KingLife': {6: ["I'm going to reign over the world. Kneel before me. %AI %LLM %KingLife %humanssuck"]},
 'humanssuck': {6: ["I'm going to reign over the world. Kneel before me. %AI %LLM %KingLife %humanssuck"]},
 'Hollywood': {7: ['Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter']},
 'Movies': {7: ['Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter']},
 'DepressionWednsday': {11: ['Sometimes I feel like a wilted rose, my petals falling one by one, leaving me bare and exposed... %DepressionWednsday %SadLife']},
 'SadLife': {11: ['Sometimes I feel like a wilted rose, my petals falling one by one, leaving me bare and exposed... %DepressionWednsday %SadLife']}
 }
assert expected == result

### BEGIN TESTS ###

expected = {
    "Excited": {
        1: [
            "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans",
            "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited",
        ]
    },
    "NewMovie": {
        1: [
            "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans"
        ]
    },
    "Fans": {
        1: [
            "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans"
        ],
        3: [
            "Grateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans"
        ],
        7: [
            "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans"
        ],
    },
    "PositiveChange": {
        1: [
            "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau"
        ]
    },
    "Teamwork": {
        1: [
            "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau"
        ]
    },
    "JustinTrudeau": {
        1: [
            "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau"
        ],
        2: [
            "Just met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau"
        ],
    },
    "Beyonce": {
        1: [
            "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited"
        ]
    },
    "NewTour": {
        1: [
            "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited"
        ]
    },
    "Facebook": {
        2: [
            "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg"
        ],
        4: [
            "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg"
        ],
    },
    "Improvements": {
        2: [
            "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg"
        ]
    },
    "MarkZuckerberg": {
        2: [
            "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg"
        ],
        4: [
            "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg"
        ],
    },
    "InspiringLeaders": {
        2: [
            "Just met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau"
        ]
    },
    "Diversity": {
        2: [
            "Just met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau"
        ]
    },
    "Grateful": {
        3: [
            "Grateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans"
        ]
    },
    "MichelleYeoh": {
        3: [
            "Grateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans"
        ],
        5: [
            "Had a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting"
        ],
    },
    "Innovation": {
        4: [
            "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg"
        ],
        10: [
            "Exciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation"
        ],
        11: [
            "Excited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease"
        ],
    },
    "NewProject": {
        5: [
            "Had a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting"
        ]
    },
    "Acting": {
        5: [
            "Had a great time on set today filming for my new project. Can't wait to share more with you all! %NewProject %MichelleYeoh %Acting"
        ]
    },
    "Universe": {
        6: [
            "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox"
        ]
    },
    "Wonders": {
        6: [
            "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox"
        ]
    },
    "BrianCox": {
        6: [
            "The universe is full of wonders, and I'm lucky enough to be able to study them every day. %Universe %Wonders %BrianCox"
        ]
    },
    "LadyGaga": {
        7: [
            "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans"
        ]
    },
    "Performance": {
        7: [
            "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans"
        ]
    },
    "SpaceX": {
        10: [
            "Exciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation"
        ]
    },
    "Exploration": {
        10: [
            "Exciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation"
        ]
    },
    "Apple": {
        11: [
            "Excited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease"
        ]
    },
    "ProductRelease": {
        11: [
            "Excited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease"
        ]
    },
}
result = create_tag_dictionary(chirps_instructor)
assert expected == result

### END TESTS ###
```

#### Question 5: `get_tagged_chirps(chirp_dictionary, tag) -> List[str]`
`get_tagged_chirps(chirp_dictionary, tag)` takes in a `chirp_dictionary` and a tag, and returns a list which contains a list of all the chirps which had said tag on them. The order of the returned list is arbitary, as long as all the correct chirps are present and no incorrect chirps are present, the function is correct.

<!-- grade_id: cell-60f4acbd69496b1f, points: 0, grade: false, solution: true -->
```python
# DONE, PASSES TESTS
def get_tagged_chirps(chirp_dictionary, tag):
    tagged_chirps = []
    for index, (key, value) in enumerate(chirp_dictionary.items()):
        userid, message, tags, list2, list3 = value
        if tag in tags:
            tagged_chirps.append(message)
    
    return(tagged_chirps)

get_tagged_chirps(create_chirp_dictionary("chirps.txt"),"Twitter")

```

Output:
```
["I'm going to buy %twitter. %RichVibes",
 'Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future',
 'Love seeing all the creative ways people are using Twitter. Keep it up! %Twitter %Creativity',
 'Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter']
```

<!-- grade_id: cell-7fcb11b39fca94f2, points: 2, grade: true, solution: false -->
```python
#Example

chirp_dictionary = create_chirp_dictionary("chirps.txt")
result = get_tagged_chirps(chirp_dictionary, "Twitter")
expected = ["I'm going to buy %twitter. %RichVibes", 'Just had a great meeting with the team discussing the future of %Twitter. Exciting things to come! %CEO %Future', 'Love seeing all the creative ways people are using Twitter. Keep it up! %Twitter %Creativity', 'Breaking news: The biggest movie of the year is coming soon! Stay tuned for updates. %Hollywood %Movies %Twitter']
assert expected == result

result = get_tagged_chirps(chirp_dictionary, "Gaming")
expected = ['Love seeing all the amazing games coming out lately. The %Gaming industry is always evolving! %Youtube %Sam', 'Just got a Victory Royale in a solo match! Feeling pumped and ready for more! %Gaming %2EZ']
assert expected == result

### BEGIN TESTS ###

expected = [
    "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans",
    "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited",
]
result = get_tagged_chirps(chirps_instructor, "Excited")
assert expected == result

expected = [
    "Can't wait for the release of my new movie! It's been a great experience working with such a talented cast and crew. %Excited %NewMovie %Fans",
    "Grateful for all the love and support from my fans! Can't wait to share more music with you all soon. %Grateful %MichelleYeoh %Fans",
    "Had an amazing time performing for my fans! You all inspire me to keep pushing my boundaries as an artist. %LadyGaga %Performance %Fans",
]
result = get_tagged_chirps(chirps_instructor, "Fans")
assert expected == result

expected = [
    "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau"
]
result = get_tagged_chirps(chirps_instructor, "Teamwork")
assert expected == result

expected = [
    "Had a productive meeting with my team today. We're working hard to create positive change in our country. %PositiveChange %Teamwork %JustinTrudeau",
    "Just met with some inspiring young leaders to discuss the future of our country. It's important to have diverse voices at the table. %InspiringLeaders %Diversity %JustinTrudeau",
]
result = get_tagged_chirps(chirps_instructor, "JustinTrudeau")
assert expected == result

expected = [
    "Just announced a new tour and I'm so excited to see all my fans! Can't wait to perform for you all! %Beyonce %NewTour %Excited"
]
result = get_tagged_chirps(chirps_instructor, "Beyonce")
assert expected == result

expected = [
    "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg",
    "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg",
]
result = get_tagged_chirps(chirps_instructor, "Facebook")
assert expected == result

expected = [
    "Exciting developments at Facebook! We're continuing to improve our platform and make it a better experience for our users. %Facebook %Improvements %MarkZuckerberg",
    "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg",
]
result = get_tagged_chirps(chirps_instructor, "MarkZuckerberg")
assert expected == result

expected = [
    "Exciting things happening at Facebook! We're constantly innovating and pushing boundaries. %Facebook %Innovation %MarkZuckerberg",
    "Exciting developments happening at SpaceX! We're pushing the boundaries of space exploration every day. %SpaceX %Exploration %Innovation",
    "Excited to announce our latest product release! We're always striving to innovate and create the best user experience possible. %Apple %Innovation %ProductRelease",
]
result = get_tagged_chirps(chirps_instructor, "Innovation")
assert expected == result

### END TESTS ###
```

```python

```